# Introduction to Quantum Mechanics

##  &#x1f3af; Objective¶
To review basic aspects of quantum mechanics.

## &#x1f4dc; Instructions
Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name, username (the prefix to your @university.ext e-mail), and student ID number in the cell below

<!-- grade_id: cell-347e783dba403114, points: 0, grade: true, solution: true -->
```python
Name = "First M. Last"
email_user_name = "username"
ID_number = 1234567

# It's useful to import these libraries. 
# You can import others or not even use these, though.
import numpy as np
import scipy
from scipy import constants
```

## Wave-Particle Duality

### &#x1f3b2; Particle-like features of light.
Which of the following phenomena are strongly associated with the particle-like nature of light. <br>
**A**. Blackbody radiation <br>
**B**. Compton Scattering  <br>
**C**. Electron Diffraction <br>
**D**. Stern-Gerlach Experiment <br>
**E**. Photoelectric effect 

<!-- grade_id: ansDuality1, points: 0, grade: false, solution: true -->
```python
# Report your answer as a list, tuple, or set containing the correct answers. 
# For example, if the answers are A and C, then ad1 = ["A", "C"]. 
# I've initialized the answer to the empty list.
ad1 = []

### BEGIN SOLUTION
ad1 = ["A", "B", "E"]
### END SOLUTION
```

<!-- grade_id: testDuality1, points: 0, grade: true, solution: false -->
```python
print("The following phenomena are associated with the particle-like nature of light:", ad1)
assert(isinstance(ad1,set) or isinstance(ad1,list) or isinstance(ad1,tuple))
assert(len(ad1) > 0)
### BEGIN HIDDEN TESTS
assert(set(map(str.casefold,ad1)) == {"a","b","e"})
### END HIDDEN TESTS
```

Output:
```
The following phenomena are associated with the particle-like nature of light: ['A', 'B', 'E']
```

<!-- grade_id: explainDuality1, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

More information about this question can be found in section 3 of the notes [From Newton to Schrodinger](https://paulwayers.github.io/IntroQChem/notes/html/History.html "See especially section 3").

=== END MARK SCHEME ===

### &#x1F5A9; Properties of photons
What is the frequency of light in Hz ($s^{-1}$) of light with wavelength 500 nm?

<!-- grade_id: ansDuality2, points: 0, grade: false, solution: true -->
```python
# Use the code box as a calculator, but report your answer as a float. 
# ansDuality2 = float. I've initialized the answer to None.
ad2 = None

### BEGIN SOLUTION
# wavelength * frequency = speed of light
ad2 = constants.c/500e-9
print("the frequency of light with wavelength 500 nm is {0:.3e} Hz".format(ad2))
### END SOLUTION
```

Output:
```
the frequency of light with wavelength 500 nm is 5.996e+14 Hz
```

<!-- grade_id: testDuality2, points: 0, grade: true, solution: false -->
```python
assert(isinstance(ad2,float))

### BEGIN HIDDEN TESTS
assert(np.isclose(ad2,constants.c/500e-9,rtol=1e-3))
### END HIDDEN TESTS
```

<!-- grade_id: explainDuality2, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

The key equations for this problem are:

$$\lambda \nu = c$$

so

$$\nu = \frac{c}{\lambda}$$

=== END MARK SCHEME ===

### &#x1fa99; Properties of photons

Doubling the wavelength of radiation doubles its frequency. (True/False)


<!-- grade_id: A_frequency_doubles, points: 0, grade: false, solution: true -->
```python
# Report your answer as a Boolean, so freq_double = True or freq_double = False
freq_double = None

### BEGIN SOLUTION
freq_double = False
### END SOLUTION
```

<!-- grade_id: T_frequency_doubles, points: 0, grade: true, solution: false -->
```python
assert(isinstance(freq_double,bool))
print("It is", freq_double, "that when the wavelength of radiation doubles its frequency does also.")
### BEGIN HIDDEN TESTS
assert(freq_double == False)
# The frequency halves, because frequency = c/wavelength
### END HIDDEN TESTS
```

Output:
```
It is False that when the wavelength of radiation doubles its frequency does also.
```

### &#x1fa99; Properties of photons

Doubling the wavelength of radiation halves its speed. (True/False)

<!-- grade_id: A_speed_halves, points: 0, grade: false, solution: true -->
```python
# Report your answer as a Boolean, so speed_halves = True or speed_halves = False
speed_halves = None

### BEGIN SOLUTION
speed_halves = False
### END SOLUTION
```

<!-- grade_id: T_speed_halves, points: 0, grade: true, solution: false -->
```python
assert(isinstance(speed_halves,bool))
print("It is", speed_halves, "that when the wavelength of radiation doubles its speed halves.")
### BEGIN HIDDEN TESTS
assert(speed_halves == False)
# The speed of light is a constant and does not depend on its wavelength.
### END HIDDEN TESTS
```

Output:
```
It is False that when the wavelength of radiation doubles its speed halves.
```

### &#x1F5A9; Properties of photons
A helium-neon laser emits light at 632.8 nm.  What is the energy of the photons generated by this laser, in Joules?

```python
# Use the code box as a calculator, but report your answer as a float. 
# E_HeNe = float. I've initialized the answer to None.
E_HeNe = None

### BEGIN SOLUTION
# E = h * frequency = h * c/wavelength
E_HeNe = constants.h * constants.c/632.8e-9
print("the energy of light with wavelength 632.8 nm is {0:.3e} J".format(E_HeNe))
### END SOLUTION
```

Output:
```
the energy of light with wavelength 632.8 nm is 3.139e-19 J
```

```python
assert(isinstance(E_HeNe,float))

### BEGIN HIDDEN TESTS
assert(np.isclose(E_HeNe,3.1391e-19,rtol=1e-3))
### END HIDDEN TESTS
```

### &#x1f3b2; Properties of photons

Which of the following changes would double the energy of a photon: <br>
**A**. Doubling its frequency <br>
**B**. Doubling its wavelength <br>
**C**. Doubling its momentum <br>
**D**. Doubling its speed <br>
**E**. Doubling its effective (relativistic) mass <br>
**F**. Doubling its wavenumber. 

<!-- grade_id: A_energy_doubles, points: 0, grade: false, solution: true -->
```python
# Report your answer as a list, tuple, or set containing the correct answers. 
# For example, if the answers are A and C, then e_doubles = ["A", "C"]. 
# I've initialized the answer to the empty list.
e_doubles = []

### BEGIN SOLUTION
e_doubles = ["A","C","E","F"]
### END SOLUTION
```

<!-- grade_id: T_energy_doubles, points: 0, grade: true, solution: false -->
```python
print("Ways you can double the energy of a photon include", e_doubles)
assert(isinstance(e_doubles,set) or isinstance(e_doubles,list) or isinstance(e_doubles,tuple))
assert(len(e_doubles) > 0)
### BEGIN HIDDEN TESTS
assert(set(map(str.casefold,e_doubles)) == {"a","c","e","f"})
### END HIDDEN TESTS
```

Output:
```
Ways you can double the energy of a photon include ['A', 'C', 'E', 'F']
```

### &#x1F5A9; Momentum from a green laser pointer
I have a high-powered green laser pointer (532 nm wavelength, 100 mW power) that I use for astronomical starspotting. If I shine this laser pointer on you, how much momentum, per second, will be transferred to you? Report your answer in SI units of kg m/s.

```python

```

<!-- grade_id: Agreenlaser, points: 0, grade: false, solution: true -->
```python
# Use the code box as a calculator, but report your answer as a float. 
# p_greenlaser = float. I've initialized the answer to None.
p_greenlaser = None

### BEGIN SOLUTION
# The energy and momentum of a single green photon is 
p_greenphoton = constants.h/532e-9  #in kg m/s
E_greenphoton = constants.h*constants.c/532e-9 #in Joules
# Based on the power, which is 100 mW = .1 J/s, we can deduce
# power = (energy of a single photon)(number of photons per second)
n_greenphotons_persecond = .1/E_greenphoton
p_greenlaser = p_greenphoton * n_greenphotons_persecond
### END SOLUTION
```

<!-- grade_id: Tgreenlaser, points: 0, grade: true, solution: false -->
```python
assert(isinstance(p_greenlaser,float))
print("the momentum transfered per second is {0:.3e} kg m/s".format(p_greenlaser))

### BEGIN HIDDEN TESTS
assert(np.isclose(p_greenlaser,3.336e-10,rtol=1e-3))
### END HIDDEN TESTS
```

Output:
```
the momentum transfered per second is 3.336e-10 kg m/s
```

### &#x1F5A9; Wavelength emitted by a radiopharmaceutical
The radioactive isotope Cobalt-60 is used in nuclear medicine to treat cancer. The energy emitted by Cobalt-60 is 1.29 x 10^11 J/mol. What is the wavelength of the emitted $\gamma$ rays?

<!-- grade_id: cell-66e8bdad8b507666, points: 0, grade: false, solution: true -->
```python
# Use the code box as a calculator, but report your answer as a float. 
# wlength_Co60 = float. I've initialized the answer to None.
wlength_Co60 = None

### BEGIN SOLUTION
# The energy is given in Joules per mole, so let's first compute the energy of a single photon,
E_photonCo60 = 1.29e11/constants.N_A
# The wavelength is then determined form E = h*frequency = hc/wavelength
wlength_C60 = constants.h * constants.c/E_photonCo60
print("the wavelength emitted by the radioactive isotope Co60 is {0:.3e} m".format(wlength_C60))
### END SOLUTION
```

Output:
```
the wavelength emitted by the radioactive isotope Co60 is 9.273e-13 m
```

```python
assert(isinstance(wlength_C60,float))

### BEGIN HIDDEN TESTS
assert(np.isclose(wlength_C60,9.273e-13,rtol=1e-3))
### END HIDDEN TESTS
```

### &#x1fa99; Davisson-Germer experiment
The Davisson-Germer experiment was among the first explicit verifications of the wave-like nature of electrons, and was foundational for modern electron diffraction methods. (True/False)

<!-- grade_id: Adualit3, points: 0, grade: false, solution: true -->
```python
# Report your answer as a Boolean, so ad3 = True or ad3 = False
ad3 = None

### BEGIN SOLUTION
ad3 = True
### END SOLUTION
```

<!-- grade_id: Tdualit3, points: 0, grade: true, solution: false -->
```python
assert(isinstance(ad3,bool))
print("The answer is:", ad3)
### BEGIN HIDDEN TESTS
assert(ad3 == True)
### END HIDDEN TESTS
```

Output:
```
The answer is: True
```

<!-- grade_id: Edualit3, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

You can find more details about the Davisson-Germer experiment in section 3.3 of the [notes on the Introduction to Quantum Mechanics](https://paulwayers.github.io/IntroQChem/notes/html/History.html).

=== END MARK SCHEME ===

### &#x1f3b2; Davisson-Germer experiment
The Davisson-Germer experiment demonstrated that if you shine a beam of electrons on a metal crystal, the result is <br>
**A**. the electrons are absorbed at “critical energies” similar to the optical (light) absorption spectrum. <br>
**B**. the electrons scatter according to the Bragg law for X-ray scattering. <br>
**C**. the electrons go right through the metal. <br>
**D**. the metal gets very hot and becomes a dull red color (stimulated blackbody emission of radiation). <br>

<!-- grade_id: Aduality3b, points: 0, grade: false, solution: true -->
```python
# Report your answer as a list, tuple, or set containing the correct answers. 
# For example, if the answers are A and C, then ad3b = ["A", "C"]. 
# I've initialized the answer to the empty list.
ad3b = []

### BEGIN SOLUTION
ad3b = ["B"]
### END SOLUTION
```

<!-- grade_id: Tduality3b, points: 0, grade: true, solution: false -->
```python
print("In the Davisson-Germer experiment", ad3b)
assert(isinstance(ad3b,set) or isinstance(ad3b,list) or isinstance(ad3b,tuple))
assert(len(ad3b) == 1)
### BEGIN HIDDEN TESTS
assert(set(map(str.casefold,ad3b)) == {"b"})
### END HIDDEN TESTS
```

Output:
```
In the Davisson-Germer experiment ['B']
```

### &#x1F5A9; Properties of photons
What is the momentum of a $\gamma$-ray photon with a wavelength of $10^{-13}$ m in SI units of ${\frac{\text{m} \cdot \text{kg}}{\text{s}}}$?

<!-- grade_id: Aduality4, points: 0, grade: false, solution: true -->
```python
# Use the code box as a calculator, but report your answer as a float. 
# ad4 = float. I've initialized the answer to None.
ad4 = None

### BEGIN SOLUTION
# momentum = h/wavelength
ad4 = constants.h/1e-13
print("the momentum of a photon with a wavelength of 1e-13 m is {0:.3e} m kg/s".format(ad4))
### END SOLUTION
```

Output:
```
the momentum of a photon with a wavelength of 1e-13 m is 6.626e-21 m kg/s
```

<!-- grade_id: Tduality4, points: 0, grade: true, solution: false -->
```python
assert(isinstance(ad4,float))

### BEGIN HIDDEN TESTS
assert(np.isclose(ad4,constants.h/1e-13,rtol=1e-3))
### END HIDDEN TESTS
```

<!-- grade_id: Eduality4, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===
The momentum of a photon can be computed from the De Broglie relation (here, better credited to Compton):

$$ p = \frac{h}{\lambda} $$

=== END MARK SCHEME ===

### &#x1F5A9; Rydberg's Law
Rydberg's law says that the wavenumber for the absorptions for a one-electron atom/ion with atomic number Z is given by the expression

$$ \tilde{\nu} = \left( 1.0974 \cdot 10^7 m^{-1}\right) Z^2 
\left( \frac{1}{n_1^2} - \frac{1}{n_2^2} \right) $$

where $1 < n_1 < n_2 < \inf$. Suppose you are given the Hydrogen atom in its ground state, $n_1=1$. What is the lowest absorption frequency?

<!-- grade_id: Aduality5, points: 0, grade: false, solution: true -->
```python
# Use the code box as a calculator, but report your answer as a float. 
# ad5 = float. I've initialized the answer to None.
ad5 = None

### BEGIN SOLUTION
wavenumber = 1.0974e7 * 1 * (1 - 1./4) #from the Rydberg formula
# frequency is speed of light times wavenumber, where wavenumber = 1/wavelength
ad5 = constants.c*wavenumber
print("the lowest absorption frequency for the hydrogen atom is {0:.3e} Hz".format(ad5))
### END SOLUTION
```

Output:
```
the lowest absorption frequency for the hydrogen atom is 2.467e+15 Hz
```

<!-- grade_id: Tduality5, points: 0, grade: true, solution: false -->
```python
assert(isinstance(ad5,float))

### BEGIN HIDDEN TESTS
assert(np.isclose(ad5,2.467e15,rtol=1e-3))
### END HIDDEN TESTS
```

<!-- grade_id: Eduality5, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

The lowest absorption frequency will correspond to exciting from the ground state to the lowest excited state, so $n_2 = 2$. Using this, we can compute the wavelength from:

$$\tilde{\nu} = (1.0974\cdot 10^7)(1^2)\left(\frac{1}{1^2} - \frac{1}{2^2} \right) $$

and then convert the wavelength to frequency using

$$ \nu = c\tilde{\nu} $$

=== END MARK SCHEME ===

### &#x1f3b2; Wave properties of particles
Which of the following experimental results are often cited as examples of the wave-likeness of particles like electrons?  
**A**. blackbody radiation  
**B**. discrete emission lines in the hydrogen spectrum  
**C**. photoelectric effect  
**D**. Compton scattering of light by a particle  
**E**. Electron scattering from a crystal

<!-- grade_id: Aduality6, points: 0, grade: false, solution: true -->
```python
# Report your answer as a list, tuple, or set containing the correct answers. 
# For example, if the answers are A and C, then ad6 = ["A", "C"]. 
# I've initialized the answer to the empty list.
ad6 = []

### BEGIN SOLUTION
ad6 = ["E"]
### END SOLUTION
```

<!-- grade_id: Tduality6, points: 0, grade: true, solution: false -->
```python
print("The following phenomena are associated with the wave-like nature of electrons:", ad6)
assert(isinstance(ad6,set) or isinstance(ad6,list) or isinstance(ad6,tuple))
assert(len(ad6) > 0)
### BEGIN HIDDEN TESTS
assert(set(map(str.casefold,ad6)) == {"e"} or set(map(str.casefold,ad6)) == {"b","e"})
### END HIDDEN TESTS
# B is a reasonable answer from the viewpoint of the Bohr model of the Hydrogen atom, but is less obvious than
# E (electron scattering).

```

Output:
```
The following phenomena are associated with the wave-like nature of electrons: ['E']
```

###  &#x1f3b2; Particle properties of waves
Which of the following experimental results are often cited as examples of the particle-likeness of radiation (light)?   
**A**. blackbody radiation <br>
**B**. discrete emission lines in the hydrogen spectrum <br>
**C**. photoelectric effect <br>
**D**. Compton scattering of light by a particle <br>
**E**. Electron scattering from a crystal <br>

<!-- grade_id: Aduality7, points: 0, grade: false, solution: true -->
```python
# Report your answer as a list, tuple, or set containing the correct answers. 
# For example, if the answers are A and C, then ad7 = ["A", "C"]. 
# I've initialized the answer to the empty list.
ad7 = []

### BEGIN SOLUTION
ad7 = ["A","C","D"]
### END SOLUTION
```

<!-- grade_id: Tdualilty7, points: 0, grade: true, solution: false -->
```python
print("The following phenomena are associated with the particle-like nature of light:", ad7)
assert(isinstance(ad7,set) or isinstance(ad7,list) or isinstance(ad7,tuple))
assert(len(ad7) > 0)
### BEGIN HIDDEN TESTS
assert(set(map(str.casefold,ad7)) == {"a","c","d"})
### END HIDDEN TESTS
```

Output:
```
The following phenomena are associated with the particle-like nature of light: ['A', 'C', 'D']
```

### &#x1F5A9; xxxProperties of photons
Suppose you are given a photon with an energy of 2 eV. What is its momentum in 
$\frac{\text{m} \cdot \text{kg}}{\text{s}}$? What is its frequency in Hz?

<!-- grade_id: Aduality9, points: 0, grade: false, solution: true -->
```python
# Use the code box as a calculator, but report your answer as float(s). 
# I've initialized the answers to None.
momentum_d9 = None
frequency_d9 = None

### BEGIN SOLUTION
# frequency = E/h
# get Planck's constant in useful units.
h_in_eVs = scipy.constants.value("Planck constant in eV/Hz")
frequency_d9 = 2.0/h_in_eVs
#Now useful to use Planck's constant in nice units.
momentum_d9 = constants.h*frequency_d9/constants.c

print("the frequency of a photon with an energy of 2 eV is {0:.3e} Hz".format(frequency_d9))
print("the momentum of a photon with an energy of 2 eV is {0:.3e} m kg/s".format(momentum_d9))
### END SOLUTION
```

Output:
```
the frequency of a photon with an energy of 2 eV is 4.836e+14 Hz
the momentum of a photon with an energy of 2 eV is 1.069e-27 m kg/s
```

<!-- grade_id: Tduality9, points: 0, grade: true, solution: false -->
```python
assert(isinstance(momentum_d9,float))
assert(isinstance(frequency_d9,float))

### BEGIN HIDDEN TESTS
assert(np.isclose(momentum_d9,1.069e-27,rtol=1e-3))
assert(np.isclose(frequency_d9,4.836e+14,rtol=1e-3))
### END HIDDEN TESTS
```

<!-- grade_id: Eduality9, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

First we can compute the frequency of the photon as:

$$ \nu = \frac{E}{h} $$

but there is the slight complication that the energy was given in electron-volts. Fortunately we have this constant built into to scipy.constants. 

The momentum of the photon can be computed from the De Broglie relation,

$$ p = \frac{h}{\lambda} = \frac{h}{\tfrac{c}{\nu}} = \frac{h \nu}{c} = \frac{E}{c} $$

Where the last formula, which was proposed long ago by Einstein and Compton and appeared in the notes, could have been used directly had you remembered it. However, because our energy is in electron-volts, it's a bit easier to use the next-to-last formula.

=== END MARK SCHEME ===

### &#x1F5A9; xxxProperties of photons
What is the momentum and energy of a photon with angular wavenumber $k=10^7 \text{m}^{-1}$?

<!-- grade_id: A_Ep_from_K, points: 0, grade: false, solution: true -->
```python
# Use the code box as a calculator, but report your answer as float(s). 
# I've initialized the answers to None.
p_from_k = None   #momentum of the photon
E_from_k = None   #Energy of the photon

### BEGIN SOLUTION
# p = h-bar * k
p_from_k = constants.hbar * 1e7
E_from_k = constants.c * p_from_k

print("the momentum of a photon with an angular wavenumber of 1e7 1/m is {0:.3e} m kg/s.".format(p_from_k))
print("the energy of a photon with an angular wavenumber of 1e7 1/m is {0:.3e} J.".format(E_from_k))
### END SOLUTION
```

Output:
```
the momentum of a photon with an angular wavenumber of 1e7 1/m is 1.055e-27 m kg/s.
the energy of a photon with an angular wavenumber of 1e7 1/m is 3.162e-19 J.
```

<!-- grade_id: T_Ep_from_k, points: 0, grade: true, solution: false -->
```python
assert(isinstance(p_from_k,float))
assert(isinstance(E_from_k,float))

### BEGIN HIDDEN TESTS
assert(np.isclose(p_from_k,1.055e-27,rtol=1e-3))
assert(np.isclose(E_from_k,3.162e-19,rtol=1e-3))
### END HIDDEN TESTS
```

<!-- grade_id: E_Ep_from_k, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

We can start with the equation for the momentum, which is easy:

$$ p = \hbar k $$


The equation for the energy can deduced directly as $E = pc$, but if you forgot this, then,

$$ E = h \nu = \frac{hc}{\lambda} = p c $$

using the Planck relation (first equality) and the De Broglie relation (last equality).

=== END MARK SCHEME ===

### &#x1F5A9; xxxDe Broglie wavelength of a baseball
During departmental baseball games, your instructor insists that the only reason he strikes out is because of the De Broglie wavelength of the baseball means that even though he swings in the correct location, he still misses. Suppose that the opposing major-league-quality hurler throws the baseball (mass = 145 g) at 100 miles-per-hour (45 m/s). What is the De Broglie wavelength of the baseball?

<!-- grade_id: A_baseball, points: 0, grade: false, solution: true -->
```python
# Use the code box as a calculator, but report your answer as float(s). 
# I've initialized the answer to None.
wl_baseball = None   #wavelength of the baseball.

### BEGIN SOLUTION
# wavelength = h/momentum = h/(mass * velocity) = h/(.145 kg * 45 m/s)
wl_baseball = constants.h/(.145*45)

print("the wavelength of the baseball is {0:.3e} m!".format(wl_baseball))
### END SOLUTION
```

Output:
```
the wavelength of the baseball is 1.015e-34 m!
```

```python
assert(isinstance(wl_baseball,float))

### BEGIN HIDDEN TESTS
assert(np.isclose(wl_baseball,1.e-34,rtol=1e-2))
### END HIDDEN TESTS
```

## The Schr&ouml;dinger Equation

### &#x270d;&#xfe0f; xxxTime-Dependent Schr&ouml;dinger Equation
What is the time-dependent Schr&ouml;dinger equation for the complex conjugate of the wavefunction, $\Psi^*$?
Put your answer in the markdown cell below. You can drag and drop an attachment (of most types) to this cell also. 

<!-- grade_id: Ase_1, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===

Taking the complex-conjugate of the time-dependent Schr&ouml;dinger equation gives:

$$ -i \hbar \frac{d \Psi^*(x,t)}{dt} = - \frac{\hbar}{2m} \frac{d^2 \Psi^*(x,t)}{dx^2} + V(x,t)\Psi^*(x,t) $$

=== END MARK SCHEME ===

### &#x1f3b2; xxxHamiltonian operator
The Hamiltonian operator corresponds to which observable property of a quantum system? <br>
**A**. Action <br>
**B**. Momentum <br>
**C**. Kinetic Energy <br>
**D**. De Broglie Wavelength <br>
**E**. Total Energy <br>
**F**. Angular Momentum <br>
**G**. Entropy  
**H**. Planck Mass

<!-- grade_id: Ase_2, points: 0, grade: false, solution: true -->
```python
# Report your answer as a list, tuple, or set containing the correct answers. 
# For example, if the answers are A and C, then ansSE2 = ["A", "C"]. 
# I've initialized the answer to the empty list.
ansSE2 = []

### BEGIN SOLUTION
ansSE2 = ["E"]
### END SOLUTION
```

<!-- grade_id: Tse_2, points: 0, grade: true, solution: false -->
```python
print("The Hamiltonian is the quantum-mechanical operator for:", ansSE2)
assert(isinstance(ansSE2,set) or isinstance(ansSE2,list) or isinstance(ansSE2,tuple))
assert(len(ansSE2) == 1)
### BEGIN HIDDEN TESTS
assert(set(map(str.casefold,ansSE2)) == {"e"})
### END HIDDEN TESTS
```

Output:
```
The Hamiltonian is the quantum-mechanical operator for: ['E']
```

## Mathematics

###  &#x1fa99; xxxMathematical Properties of the wavefunction
A probability density can be negative. (True/False)

<!-- grade_id: AProb_canbe_negative, points: 0, grade: false, solution: true -->
```python
prob_canbe_negative = None

### BEGIN SOLUTION
prob_canbe_negative = False
### END SOLUTION
```

<!-- grade_id: TProb_canbe_negative, points: 0, grade: true, solution: false -->
```python
assert(isinstance(prob_canbe_negative,bool))
print("It is", prob_canbe_negative, "that a probability density can be negative.")
### BEGIN HIDDEN TESTS
assert(prob_canbe_negative == False)
### END HIDDEN TESTS
```

Output:
```
It is False that a probability density can be negative.
```

### &#x1f3b2; xxxComplex Conjugation
Let $z$ be a complex number. If $w$ is the product of $z$ and its complex 
conjugate, $w = z z^*$, which of the following is **always** true about $w$: <br>
**A**. w is an imaginary number. <br>
**B**. w is a complex number. <br>
**C**. w is nonzero real number. <br>
**D**. w is a nonnegative real number. <br>
**E**. w is a nonzero complex number. <br>
**F**. None of the above

<!-- grade_id: Azzstar, points: 0, grade: false, solution: true -->
```python
# Report your answer as a list, tuple, or set containing the correct answers. 
# For example, if the answers are A and C, then zzstar = ["A", "C"]. 
# I've initialized the answer to the empty list.
zzstar = []

### BEGIN SOLUTION
zzstar = ["B","D"]
### END SOLUTION
```

<!-- grade_id: Tzzstar, points: 0, grade: true, solution: false -->
```python
print("The product of a number and its complex conjugate is always", zzstar)
assert(isinstance(zzstar,set) or isinstance(zzstar,list) or isinstance(zzstar,tuple))
assert(len(zzstar) > 0)
### BEGIN HIDDEN TESTS
assert(set(map(str.casefold,zzstar)) == {"b","d"})
### END HIDDEN TESTS
```

Output:
```
The product of a number and its complex conjugate is always ['B', 'D']
```

### &#x270d;&#xfe0f; Complex Conjugation
What is the complex conjugate of 

$$ \Psi(x,t) = A e^{(a+bi)(kx - \omega t)} $$


<!-- grade_id: aMath1, points: 0, grade: true, solution: true -->
=== BEGIN MARK SCHEME ===
The complex conjugate is obtained by replacing $i$ with $-i$. So

$$ \Psi^*(x,t) = A e^{(a-bi)(kx - \omega t)} $$

I would accept an answer where it was not assumed that the constants in the expression were real, e.g.,

$$ \Psi^*(x,t) = A^* e^{(a^*-b^*i)(k^*x - \omega^* t)} $$

=== END MARK SCHEME ===

### &#x1fa99; Eigenfunctions of the kinetic-energy operator
*Every* eigenfunction of the momentum operator is also an eigenfunction of the kinetic-energy operator. (True/False)

<!-- grade_id: A_ke_eigenfunction, points: 0, grade: false, solution: true -->
```python
# Report your answer as a Boolean, so is_also_eigenfunction = True or = False
is_also_eigenfunction = None

### BEGIN SOLUTION
is_also_eigenfunction = True
### END SOLUTION
```

```python
assert(isinstance(is_also_eigenfunction,bool))
print("The answer is:", is_also_eigenfunction)
### BEGIN HIDDEN TESTS
assert(is_also_eigenfunction == True)
### END HIDDEN TESTS
```

Output:
```
The answer is: True
```


$$\hat{p} \psi(x) = \lambda \psi(x) $$

$$ \hat{T} = \frac{\hat{p}^2}{2m} $$

$$ \hat{T} \psi(x) = \tfrac{1}{2m} \hat{p} \hat{p} \psi(x) =  \tfrac{1}{2m} \hat{p} \lambda \psi(x) =  \tfrac{1}{2m} \lambda^2 \psi(x) $$

The reverse is also true, but it's more subtle. You can use the fact that $\hat{p} = 2m \sqrt{\hat{T}}$, but this is not quite true; $\cos a x$ and $sin a x$ are eigenfunctions of the kinetic energy but not the momentum. The general result is that, given an operator, $\hat{Q}$, with has eigenfunctions

$$ \hat{Q} \psi_k(x) = \theta_k \psi_k(x) $$

then any (analytic) function of $\hat{Q}$, has the same eigenfunctions, and the values are:

$$ f(\hat{Q}) \psi_k(x) = f(\theta_k) \psi_k(x) $$


&#x1f4dd; 
&#x1f500;

```python

```